Faqatgina `YOUR CODE HERE` yoki "YOUR ANSWER HERE" deyilgan joyga dasturingizni yoki javobingizni yozing

```python
NAME = ""
COLLABORATORS = "Mohirdev"
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
# Ma'lumotlar tahlili

## 2.1.1. N-O'lchovli massivlarga ishlov berish


```python
# Numpy kutubxonasini chaqirib olish
import numpy as np
```

### Masala(1): 1-o'lchamli massivlar yaratish (4 points)

**Masala sharti :** Elementlari ixtiyoriy  *n* ta  (*n* > 0, butun son) bo'lgan 1-o'lchamli massiv (array) yaratuvchi funksiyani tuzing. 

**Yordam :**  `array` yoki `random` funksiyasidan foydalaning. 


<!-- grade_id: one_dim, points: 0, grade: false, solution: true -->
```python
def one_dim(n):
    """ n ta elementdan tashkil topgan 1-o'lchamli massiv   
    """
    arr = np.arange(0,n)
    # print(arr)
    return arr
one_dim(n)
    
    # raise NotImplementedError()
        
        
    
```

Sizning funksiyangiz `n`elementlari soni, `ndim` massiv o'lchami, hamda numpy array ekanligi orqali tekshiriladi: funksiyangizni pastdagi cellda tekshirib ko'ring

<!-- grade_id: correct_one_dim, points: 4, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
for i in range(1, 10):
    assert len(one_dim(i)) == i  # n elementlar soni
    arr = one_dim(i) # massiv
    assert arr.ndim == 1 # massiv o'lchami
    assert type(arr) == np.ndarray # #numpy array 
```

Output:
```
<ipython-input-8-dd14b0d4ad9b> in <module>
      1 """Check that squares returns the correct output for several inputs"""
      2 for i in range(1, 10):
----> 3     assert len(one_dim(i)) == i  # n elementlar soni
      4     arr = one_dim(i) # massiv
      5     assert arr.ndim == 1 # massiv o'lchami

<ipython-input-7-03168514523d> in one_dim(n)
      2     """ n ta elementdan tashkil topgan 1-o'lchamli massiv   
      3     """
----> 4     arr = np.arange(0,n)
      5     # print(arr)
      6     return arr

NameError: name 'np' is not defined
```

### Masala(2) 2-o'lchamli massivlar yaratish (6 points)

**Masala :** Shakli `shape(r, c)` bo'lgan 2-o'lchamli massiv (array) yaratuvchi funksiyani tuzing.  

**Yordam :** numpy kutubxonasida massivlar yaratuvchi funksiyalarining istalganidan foydlanishingiz mumkin. 


<!-- grade_id: two_dim, points: 0, grade: false, solution: true -->
```python
def two_dim(r, c):
    """ r ta qator, c ta ustundan tashkil topgan 2-o'lchamli massiv  
    """
    arr2 = np.random.randint(10, size= (r, c))
    return arr2
two_dim(r, c)
     # raise NotImplementedError()
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-11-c6114d2c517f> in <module>
      5     arr2
      6     return arr2
----> 7 two_dim(r, c)
      8      # raise NotImplementedError()

NameError: name 'r' is not defined
```

Funksiyangizni pastdagi cellda test qilib ko'ring

<!-- grade_id: correct_two_dim, points: 6, grade: true, solution: false -->
```python
arr = two_dim(3, 2) # massiv
assert arr.shape == (3, 2) # shape
assert arr.ndim == 2 # massiv o'lchami
assert type(arr) == np.ndarray # #numpy array 
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-10-5b09d6b5ed1f> in <module>
----> 1 arr = two_dim(3, 2) # massiv
      2 assert arr.shape == (3, 2) # shape
      3 assert arr.ndim == 2 # massiv o'lchami
      4 assert type(arr) == np.ndarray # #numpy array

NameError: name 'two_dim' is not defined
```

```python

```