Faqatgina `YOUR CODE HERE` yoki "YOUR ANSWER HERE" deyilgan joyga dasturingizni yoki javobingizni yozing

```python
NAME = ""
COLLABORATORS = "Mohirdev"
```

---

---
# Ma'lumotlar tahlili

## 2.1.1. N-O'lchovli massivlarga ishlov berish

## Vazifa-2

```python
# Numpy kutubxonasini chaqirib olish
import numpy as np
```

### Masala(1): 2-o'lchamli massiv yaratish (1 point)

**Masala sharti :** Elementlari 0 dan 9 gacha (9 ning o'zi massiv elementiga kirmaydi), qadami esa 1 ga teng bo'lgan (3, 3) o'lchamli massiv yarating. 

**Yordam :**  `arange`dan foydalanish ishni osonlashtiradi 


<!-- grade_id: 3x3array, points: 0, grade: false, solution: true -->
```python
def nd_array():
    """Elementlari 0 dan 9 gacha (9 ning o'zi massiv elementiga kirmaydi), qadami esa 1 ga teng bo'lgan (3, 3) o'lchamli massiv """
    arr = np.arange(9)
    return arr.reshape(3,3)
nd_array()
    #raise NotImplementedError()
```

Sizning funksiyangiz massiv elementlari soni, `shape` massiv shakli, hamda elementlari 0~10 oralig'idagi numpy array ekanligi orqali tekshiriladi. Bundan tashqari bir qancha maxfiy testlardan o'tkaziladi (maslahatimiz aldashga urinmang). 

<!-- grade_id: correct_3x3array, points: 1, grade: true, solution: false -->
```python
"""Funksiyani tekshirish"""
arr = nd_array() # massiv
assert arr.size == 9  # elementlar soni
assert type(arr) == np.ndarray # #numpy array 
```

###  Masala(2): Indeks yordamida 2-o'lchamli massivlardan elementlarni  kesib olish (3 points)

**Masala :** Yuqorida yaratilgan massivning elementlari 6 va 7 ga teng bo'lgan qismini kesib olish funksiyasini tuzing.

**Yordam :** 6 va 7 turgan indekslarni aniqlab olib, shu indekslar yordamida kesib oling. **Natija:** array([6, 7]) ko'rinishida bo'lishi kerak. 

<!-- grade_id: 2dsliced_array, points: 0, grade: false, solution: true -->
```python
def sliced_array_2d():
    """Yuqorida yaratilgan massivning elementlari 6 va 7 ga teng bo'lgan qismini kesib olish funksiyasi """
    sl_arr = arr[2, :2]
    return sl_arr
sliced_array_2d()
    #raise NotImplementedError()
```

Sizning funksiyangiz massiv elementlari soni, `shape` massiv shakli, hamda elementlari numpy array ekanligi orqali tekshiriladi. Bundan tashqari bir qancha maxfiy testlardan o'tkaziladi (maslahatimiz aldashga urinmang). 

<!-- grade_id: correct_2dsliced_array, points: 3, grade: true, solution: false -->
```python
"""Funksiyani tekshirish"""
sl_arr_2d = sliced_array_2d() # kesib olingan massiv
assert sl_arr_2d.size == 2  # elementlar soni
assert type(sl_arr_2d) == np.ndarray # #numpy array 
```

###  Masala(3): Indeks yordamida 3-o'lchamli massivdan elementlarni kesib olish (4 points)

**Masala :**  3-o'lchamli massivni quyidagi listdan yarating, hamda 12, 13, 15, va 16  elementlarini kesib oling. 

                  [[[ 0,  1,  2],
                    [ 3,  4,  5],
                    [ 6,  7,  8]],

                   [[ 9, 10, 11],
                    [12, 13, 14],
                    [15, 16, 17]],

                   [[18, 19, 20],
                    [21, 22, 23],
                    [24, 25, 26]]]

**Yordam :** Avval massivni ikki o'lchamli massivga keltirib, keyin kesib olishni amalga oshiring. 


**Natija:**  

             array([[12, 13],
                    [15, 16]])
ko'rinishida bo'lishi kerak. 

Sizning funksiyangiz massiv elementlari soni, `shape` massiv shakli, hamda elementlari numpy array ekanligi orqali tekshiriladi. Bundan tashqari bir qancha maxfiy testlardan o'tkaziladi (maslahatimiz aldashga urinmang). 

<!-- grade_id: 3dsliced_arr, points: 0, grade: false, solution: true -->
```python
def sliced_array_3d():
    """3-o'lchamli massivni yuqoridagi listdan yaratish, hamda 12, 13, 15, va 16 elementlarini kesib olish funksiyasi"""
    arr3 = np.array([[[0,1,2],
                 [3,4,5],
                 [6,7,8]],
                
                [[9,10,11],
                 [12,13,14],
                 [15,16,17]],
                
                [[18,19,20],
                 [21,22,23],
                 [24,25,26]]])
    sl_arr = arr3[1][1:3, 0:2]
    return sl_arr
sliced_array_3d
   


    #raise NotImplementedError()
```

<!-- grade_id: correct_3dsliced_arr, points: 4, grade: true, solution: false -->
```python
"""Funksiyani tekshirish"""
sl_arr_3d = sliced_array_3d() # kesib olingan massiv
assert sl_arr_3d.size == 4  # elementlar soni
assert type(sl_arr_3d) == np.ndarray # #numpy array 
```

###  Masala(4): Boolean indeks yordamida 2-o'lchamli massivlardan elementlarni  kesib olish (2 points)

**Masala :** Bizlarda,  
        
        ['Hasan', 'Husan', 'Javohir', 'Elyor', 'Hasan', 'Javohir', 'Elyor']

**list**dagi kishilar va ularning **ma'lumot**lari 
       
      [[5, 6, 1, 1],
       [9, 1, 1, 1],
       [7, 7, 4, 2],
       [1, 5, 1, 9],
       [9, 9, 4, 5],
       [7, 5, 9, 6],
       [5, 3, 7, 4]]
ko'rinishda keltirilgan. 


Mana shu ma'lumotlardan **Javohir** va **Elyor** ismlariga tegishli ma'lumotlarni Boolean indekslash yordamida kesib olish funksiyasini tuzing. 
**Yordam :** Boolean indekslash video darsligiga murojat qiling. **Natija:** Natija (4, 4) o'lchamli massiv ko'rinishida, hamda dastlabki ikkita qattori quyidagicha

    array([[7, 7, 4, 2],
           [1, 5, 1, 9],
           ...
           ...        ]])

<!-- grade_id: boolean_idx, points: 0, grade: false, solution: true -->
```python
def boolean_slicing():
    """Boolean indekslash yordamida ma'lumotlarni kesib olish """
    names = np.array(['Hasan', 'Husan', 'Javohir', 'Elyor', 'Hasan', 'Javohir', 'Elyor'])
    data = np.array([[5, 6, 1, 1],
   [9, 1, 1, 1],
   [7, 7, 4, 2],
   [1, 5, 1, 9],
   [9, 9, 4, 5],
   [7, 5, 9, 6],
   [5, 3, 7, 4]])
    mask = (names == 'Javohir') | (names == 'Elyor')
    return data[mask]
    
    #raise NotImplementedError()
```

Sizning funksiyangiz massiv elementlari soni, `shape` massiv shakli, hamda elementlari numpy array ekanligi orqali tekshiriladi. Bundan tashqari bir qancha maxfiy testlardan o'tkaziladi (maslahatimiz aldashga urinmang). 

<!-- grade_id: correct_boolean_idx, points: 2, grade: true, solution: false -->
```python

"""Funksiyani tekshirish"""
bool_idx = boolean_slicing() # kesib olingan massiv
assert bool_idx.size == 16  # elementlar soni
assert type(bool_idx) == np.ndarray # #numpy array 
```